---
title: 入門
description: 在入門章節中我們介紹 Python 的資料科學開發環境、基礎語法與基本的變數型別，透過本章節的影片說明與互動練習，您將暸解 Google Colaboratory、Anaconda、Jupyter Notebook/Lab、賦值、註解、數值、文字與布林。
---

## 賦值

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
```

Pytho 使用 `=` 來進行賦値。

`@instructions`
- 將 `87` 指派給 `stupid_int` 物件

`@hint`
- `stupid_int = 87` 

`@pre_exercise_code`

```{python}
# Load datasets and packages here.
```

`@sample_code`

```{python}
# 將 87 指派給 stupid_int
stupid_int = ___
```

`@solution`

```{python}
# 將 87 指派給 stupid_int
stupid_int = 87
```

`@sct`

```{python}
# Update this to something more informative.
success_msg("太棒了！")
```

## 註解

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
```

Python 使用 `#` 來做註解。

`@instructions`
- 在程式碼中透過 `#` 來做註解

`@hint`
- 直接按送出。

`@pre_exercise_code`

```{python}
# Load datasets and packages here.
```

`@sample_code`

```{python}
# 在 # 之後撰寫註解
```

`@solution`

```{python}
# 在 # 之後撰寫註解
```

`@sct`

```{python}
# Update this to something more informative.
success_msg("太棒了！")
```
